---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

这是一个例子，主要是为了提供这种交互页面给大家学习。
例子都是没有意义，主要为了展示。

`@instructions`
- 输入`1+1` 查看答案
- 输入`1+2` 查看答案
- 输入`1+2` 查看答案
- 输入`1+2` 查看答案

`@hint`
没有什么提示，没有什么提示，没有什么提示，没有什么提示

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
